```python
# Initialize Otter
import otter
grader = otter.Notebook("lab_ass03_2024.ipynb")
```

# Computer Assignment 3

Make sure the code below works and plays audio that you can hear.

```python
import numpy as np
import sounddevice as sd

fs = 8192
y = np.loadtxt('/Users/nicolls/Desktop/handel8192.txt')
sd.play(y,fs)
```

***
**Task 1**:  FIR filtering of narrowband interference
***

We synthesise a signal corrupted by narrowband interference at frequency $\omega_I = 2 \pi/10 = 8 \pi/40$ rad/s.

```python
yc = y + 1.0*np.sin(2*np.pi/10*np.arange(len(y)))
```

Use a filter design tool with the "equiripple" Parks–McClellan algorithm to find a minimum order finite impulse response notch or bandstop filter to null the interference.  The filter should pass all frequencies $\omega < 6 \pi/40$ and $\omega > 10\pi/40$ with less than 0.1dB of ripple, and it should provide at least 60dB of attenuation over the range $7\pi/40 < \omega < 9\pi/40$.  Place the impulse response values in the vector `bv` below, and confirm that the filter evidently eliminates the interference (and some of the signal) from the corrupted audio clip.

If you're using pyFDA then it's possible to copy the filter coefficients to the clipboard from the "b,a" tab.  Open the "CSV" settings dialog, choose "Table mode" as "Rows", and select the "Clipboard" output.  You can then copy to clipboard, paste the values into a Jupyter cell, and add some minor edits to generate the required vector.

```python
from scipy import signal
import matplotlib.pyplot as plt

# Null (identity) filter below
a1v = np.ones((1,))  # no poles
b1v = np.ones((1,))  # single impulse

# Update FIR filter impulse response using results of filter design process
b1v = ...

# Plot frequency response corresponding to bv, av above
(w1v,H1v) = signal.freqz(b1v,a1v)

# Plot frequency response magnitude and phase
fig, axs = plt.subplots(2, 1, constrained_layout=True)

axs[0].plot(w1v, np.abs(H1v))
axs[0].set_title('Filter frequency response')
axs[0].set_ylabel('$|H(e^{j \omega})|$')

axs[1].plot(w1v, np.angle(H1v))
axs[1].set_xlabel('$\omega$')
axs[1].set_ylabel('$\t{angle}~H(e^{j \omega})$');
```

```python
# Filter signal (numerator is bv and usually av[0]=1)
y1v = signal.lfilter(b1v,a1v,yc)
#sd.play(y1v,fs)
```

```python
grader.check("lab_ass03_q1a")
```

***
**Task 2**:  IIR filtering of narrowband interference
***

We synthesise a signal corrupted by narrowband interference at frequency $\omega_I = 2 \pi/10 = 8 \pi/40$ rad/s.

```python
yc = y + 1.0*np.sin(2*np.pi/10*np.arange(len(y)))
```

Design an infinite impulse response filter that satisfies the same requirements as above using the "Elliptic" method.  You should observe that the IIR filter can be implemented with significantly less computation than the FIR filter, but it will have a nonlinear phase response.  You can confirm that this nonlinear phase characteristic does not appear to cause audible distortions in the output.  Apparently our ears are not sensitive to phase.

```python
from scipy import signal
import matplotlib.pyplot as plt

# Null (identity) filter below
a2v = np.ones((1,))  # no poles
b2v = np.ones((1,))  # single impulse

# Update filter parameters using results from filter design tool
b2v = ...
a2v = ...

# Plot frequency response corresponding to bv, av above
(w2v,H2v) = signal.freqz(b2v,a2v)

# Plot frequency response magnitude and phase
fig, axs = plt.subplots(2, 1, constrained_layout=True)

axs[0].plot(w2v, np.abs(H2v))
axs[0].set_title('Filter frequency response')
axs[0].set_ylabel('$|H(e^{j \omega})|$')

axs[1].plot(w2v, np.angle(H2v))
axs[1].set_xlabel('$\omega$')
axs[1].set_ylabel('$\t{angle}~H(e^{j \omega})$');
```

```python
# Filter signal (numerator is bv and usually av[0]=1)
y2v = signal.lfilter(b2v,a2v,yc)
#sd.play(y2v,fs)
```

```python
grader.check("lab_ass03_q1b")
```

***
**Task 3**:  Upsampling by rational factor
***

Suppose we want to do sampling rate conversion of the signal `y` sampled at frequency `fs` to produce `yu` that can be played using `sd.play(yu3d2,fs*3/2)`.  This requires expanding the signal by a factor of 3, filtering the result with a lowpass filter with cutoff $\omega_c = \pi/3$, and discarding every second sample from the result.

Complete the function `sigu3d2(y)` below that takes the input signal `y` and generates the required upsampled signal `yu`.  You can test that your function is working using `sd.play`.

Your lowpass filter should have less than 1dB of ripple for $\omega < 0.8 \omega_c$ and at least 40dB of attenuation for $\omega > 1.2 \omega_c$.  You can use any filter structure you like, but try to see if you can keep computational requirements low.

```python
def sigu3d2(y):

  # Lowpass filter with cutoff wc=pi/3
  av = ...
  bv = ...
  bv = 3*bv  # need additional filter gain 

  # Input expanded by factor 3
  yu3 = np.zeros((3*len(y),))
  yu3[::3] = y;

  # Interpolation using lowpass filter
  yu3f = ...

  # Result downsampled by factor 2
  yu = yu3f[::2]
  return yu
```

```python
# Upsample and play audio
yu = sigu3d2(y)
sd.play(yu,fs*3/2)
```

```python
grader.check("lab_ass03_q1c")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for lab_ass03_q1a (points: None)
```
>>> assert (len(a1v) == 1) & (a1v[0] == 1)
```
failure_message: need len(a1v)==1 and a1v[0]==1
success_message: have len(a1v)==1 and a1v[0]==1
```
>>> assert b1v.ndim == 1
```
failure_message: b1v does not have correct dimension of one
success_message: b1v has correct dimension of one


Tests for lab_ass03_q1b (points: None)
```
>>> assert len(a2v) != 1
```
failure_message: for IIR filter we do not have len(a2v)!=1 as required
success_message: for IIR filter we have len(a2v)!=1 as required
```
>>> assert (b2v.ndim == 1) & (a2v.ndim == 1)
```
failure_message: b2v and a2v do not have correct dimension of one
success_message: b2v and a2v have correct dimension of one


Tests for lab_ass03_q1c (points: None)
```
>>> yu = sigu3d2(y)
>>> assert yu.ndim == 1
```
failure_message: upsampled output does not have required dimension
success_message: upsampled output has required dimension
```
>>> yu = sigu3d2(y)
>>> assert (len(yu) / len(y) > 1.3) & (len(yu) / len(y) < 1.7)
```
failure_message: upsampled signal length not ~1.5x input signal length
success_message: upsampled signal length ~1.5x input signal length
